Cross-modal interaction of Alpha Activity does not reflect inhibition of early sensory processing: A frequency tagging study using EEG and MEG

Author response:

The following is the authors’ response to the current reviews.

I thank the authors for their clarifications. The manuscript is much improved now, in my opinion. The new power spectral density plots and revised Figure 1 are much appreciated. However, there is one remaining point that I am unclear about. In the rebuttal, the authors state the following: "To directly address the question of whether the auditory signal was distracting, we conducted a follow-up MEG experiment. In this study, we observed a significant reduction in visual accuracy during the second block when the distractor was present (see Fig. 7B and Suppl. Fig. 1B), providing clear evidence of a distractor cost under conditions where performance was not saturated."

I am very confused by this statement, because both Fig. 7B and Suppl. Fig. 1B show that the visual- (i.e., visual target presented alone) has a lower accuracy and longer reaction time than visual+ (i.e., visual target presented with distractor). In fact, Suppl. Fig. 1B legend states the following: "accuracy: auditory- - auditory+: M = 7.2 %; SD = 7.5; p = .001; t(25) = 4.9; visual- - visual+: M = -7.6%; SD = 10.80; p < .01; t(25) = -3.59; Reaction time: auditory- - auditory +: M = -20.64 ms; SD = 57.6; n.s.: p = .08; t(25) = -1.83; visual- - visual+: M = 60.1 ms ; SD = 58.52; p < .001; t(25) = 5.23)."

These statements appear to directly contradict each other. I appreciate that the difficulty of auditory and visual trials in block 2 of MEG experiments are matched, but this does not address the question of whether the distractor was actually distracting (and thus needed to be inhibited by occipital alpha). Please clarify.

We apologize for mixing up the visual and auditory distractor cost in our rebuttal. The reviewer is right in that our two statements contradict each other.

To clarify: In the EEG experiment, we see significant distractor cost for auditory distractors in the accuracy (which can be seen in SUPPL Fig. 1A). We also see a faster reaction time with auditory distractors, which may speak to intersensory facilitation. As we used the same distractors for both experiments, it can be assumed that they were distracting in both experiments.

In our follow-up MEG-experiment, as the reviewer stated, performance in block 2 was higher than in block 1, even though there were distractors present. In this experiment, distractor cost and learning effects are difficult to disentangle. It is possible that participants improved over time for the visual discrimination task in Block 1, as performance at the beginning was quite low. To illustrate this, we divided the trials of each condition into bins of 10 and plotted the mean accuracy in these bins over time (see Author response image 1). Here it can be seen that in Block 2, there is a more or less stable performance over time with a variation < 10 %. In Block 1, both for visual as well as auditory trials, an improvement over time can be seen. This is especially strong for visual trials, which span a difference of > 20%. Note that the mean performance for the 80-90 trial bin was higher than any mean performance observed in Block 2.

Additionally, the same paradigm has been applied in previous investigations, which also found distractor costs for the here-used auditory stimuli in blocked and non-blocked designs. See:

Mazaheri, A., van Schouwenburg, M. R., Dimitrijevic, A., Denys, D., Cools, R., & Jensen, O. (2014). Region-specific modulations in oscillatory alpha activity serve to facilitate processing in the visual and auditory modalities. NeuroImage, 87, 356–362. https://doi.org/10.1016/j.neuroimage.2013.10.052

Van Diepen, R & Mazaheri, A 2017, 'Cross-sensory modulation of alpha oscillatory activity: suppression, idling and default resource allocation', European Journal of Neuroscience, vol. 45, no. 11, pp. 1431-1438. https://doi.org/10.1111/ejn.13570

Accuracy development over time in the MEG experiment.

During block 1, a performance increase over time can be observed for visual as well as for auditory stimuli. During Block 2, performance is stable over time. Data are presented as mean ± SEM. N = 27 (one participant was excluded from this analysis, as their trial count in at least one condition was below 90 trials).

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

In this study, Brickwedde et al. leveraged a cross-modal task where visual cues indicated whether upcoming targets required visual or auditory discrimination. Visual and auditory targets were paired with auditory and visual distractors, respectively. The authors found that during the cue-to-target interval, posterior alpha activity increased along with auditory and visual frequency-tagged activity when subjects were anticipating auditory targets. The authors conclude that their results disprove the alpha inhibition hypothesis, and instead implies that alpha "regulates downstream information transfer." However, as I detail below, I do not think the presented data irrefutably disproves the alpha inhibition hypothesis. Moreover, the evidence for the alternative hypothesis of alpha as an orchestrator for downstream signal transmission is weak. Their data serves to refute only the most extreme and physiologically implausible version of the alpha inhibition hypothesis, which assumes that alpha completely disengages the entire brain area, inhibiting all neuronal activity.

We thank the reviewer for taking the time to provide additional feedback and suggestions and we improved our manuscript accordingly.

(1) Authors assign specific meanings to specific frequencies (8-12 Hz alpha, 4 Hz intermodulation frequency, 36 Hz visual tagging activity, 40 Hz auditory tagging activity), but the results show that spectral power increases in all of these frequencies towards the end of the cue-to-target interval. This result is consistent with a broadband increase, which could simply be due to additional attention required when anticipating auditory target (since behavioral performance was lower with auditory targets, we can say auditory discrimination was more difficult). To rule this out, authors will need to show a power spectral density curve with specific increases around each frequency band of interest. In addition, it would be more convincing if there was a bump in the alpha band, and distinct bumps for 4 vs 36 vs 40 Hz band.

This is an interesting point with several aspects, which we will address separately

Broadband Increase vs. Frequency-Specific Effects:

The suggestion that the observed spectral power increases may reflect a broadband effect rather than frequency-specific tagging is important. However, Supplementary Figure 11 shows no difference between expecting an auditory or visual target at 44 Hz. This demonstrates that (1) there is no uniform increase across all frequencies, and (2) the separation between our stimulation frequencies was sufficient to allow differentiation using our method.

Task Difficulty and Performance Differences:

The reviewer suggests that the observed effects may be due to differences in task difficulty, citing lower performance when anticipating auditory targets in the EEG study. This issue was explicitly addressed in our follow-up MEG study, where stimulus difficulty was calibrated. In the second block—used for analysis—accuracy between auditory and visual targets was matched (see Fig. 7B). The replication of our findings under these controlled conditions directly rules out task difficulty as the sole explanation. This point is clearly presented in the manuscript.

Power Spectrum Analysis:

The reviewer’s suggestion that our analysis lacks evidence of frequency-specific effects is addressed directly in the manuscript. While we initially used the Hilbert method to track the time course of power fluctuations, we also included spectral analyses to confirm distinct peaks at the stimulation frequencies. Specifically, when averaging over the alpha cluster, we observed a significant difference at 10 Hz between auditory and visual target expectation, with no significant differences at 36 or 40 Hz in that cluster. Conversely, in the sensor cluster showing significant 36 Hz activity, alpha power did not differ, but both 36 Hz and 40 Hz tagging frequencies showed significant effects These findings clearly demonstrate frequency-specific modulation and are already presented in the manuscript.

(2) For visual target discrimination, behavioral performance with and without the distractor is not statistically different. Moreover, the reaction time is faster with distractor. Is there any evidence that the added auditory signal was actually distracting?

We appreciate the reviewer’s observation regarding the lack of a statistically significant difference in behavioral performance for visual target discrimination with and without the auditory distractor. While this was indeed the case in our EEG experiment, we believe the absence of an accuracy effect may be attributable to a ceiling effect, as overall visual performance approached 100%. This high baseline likely masked any subtle influence of the distractor.

To directly address the question of whether the auditory signal was distracting, we conducted a follow-up MEG experiment. In this study, we observed a significant reduction in visual accuracy during the second block when the distractor was present (see Fig. 7B and Suppl. Fig. 1B), providing clear evidence of a distractor cost under conditions where performance was not saturated.

Regarding the faster reaction times observed in the presence of the auditory distractor, this phenomenon is consistent with prior findings on intersensory facilitation. Auditory stimuli, which are processed more rapidly than visual stimuli, can enhance response speed to visual targets—even when the auditory input is non-informative or nominally distracting (Nickerson, 1973; Diederich & Colonius, 2008; Salagovic & Leonard, 2021). Thus, while the auditory signal may facilitate motor responses, it can simultaneously impair perceptual accuracy, depending on task demands and baseline performance levels.

Taken together, our data suggest that the auditory signal does exert a distracting influence, particularly under conditions where visual performance is not at ceiling. The dual effect—facilitated reaction time but reduced accuracy—highlights the complexity of multisensory interactions and underscores the importance of considering both behavioral and neurophysiological measures.

(3) It is possible that alpha does suppress task-irrelevant stimuli, but only when it is distracting. In other words, perhaps alpha only suppresses distractors that are presented simultaneously with the target. Since the authors did not test this, they cannot irrefutably reject the alpha inhibition hypothesis.

The reviewer’s claim that we did not test whether alpha suppresses distractors presented simultaneously with the target is incorrect. As stated in the manuscript and supported by our data (see point 2), auditory distractors were indeed presented concurrently with visual targets, and they were demonstrably distracting. Therefore, the scenario the reviewer suggests was not only tested—it forms a core part of our design.

Furthermore, it was never our intention to irrefutably reject the alpha inhibition hypothesis. Rather, our aim was to revise and expand it. If our phrasing implied otherwise, we have now clarified this in the manuscript. Specifically, we propose that alpha oscillations:

(a) Exhibit cyclic inhibitory and excitatory dynamics;

(b) Regulate processing by modulating transfer pathways, which can result in either inhibition or facilitation depending on the network context.

In our study, we did not observe suppression of distractor transfer, likely due to the engagement of a supramodal system that enhances both auditory and visual excitability. This interpretation is supported by prior findings (e.g., Jacoby et al., 2012), which show increased visual SSEPs under auditory task load, and by Zhigalov et al. (2020), who found no trial-by-trial correlation between alpha power and visual tagging in early visual areas, despite a general association with attention.

Recent evidence (Clausner et al., 2024; Yang et al., 2024) further supports the notion that alpha oscillations serve multiple functional roles depending on the network involved. These roles include intra- and inter-cortical signal transmission, distractor inhibition, and enhancement of downstream processing (Scheeringa et al., 2012; Bastos et al., 2015; Zumer et al., 2014). We believe the most plausible account is that alpha oscillations support both functions, depending on context.

To reflect this more clearly, we have updated Figure 1 to present a broader signal-transfer framework for alpha oscillations, beyond the specific scenario tested in this study.

We have now revised Figure 1 and several sentences in the introduction and discussion, to clarify this argument.

L35-37: Previous research gave rise to the prominent alpha inhibition hypothesis, which suggests that oscillatory activity in the alpha range (~10 Hz) plays a mechanistic role in selective attention through functional inhibition of irrelevant cortical areas (see Fig. 1; Foxe et al., 1998; Jensen & Mazaheri, 2010; Klimesch et al., 2007).

L60-65: In contrast, we propose that functional and inhibitory effects of alpha modulation, such as distractor inhibition, are exhibited through blocking or facilitating signal transmission to higher order areas (Peylo et al., 2021; Yang et al., 2023; Zhigalov & Jensen, 2020; Zumer et al., 2014), gating feedforward or feedback communication between sensory areas (see Fig. 1; Bauer et al., 2020; Haegens et al., 2015; Uemura et al., 2021).

L482-485: This suggests that responsiveness of the visual stream was not inhibited when attention was directed to auditory processing and was not inhibited by occipital alpha activity, which directly contradicts the proposed mechanism behind the alpha inhibition hypothesis.

L517-519: Top-down cued changes in alpha power have now been widely viewed to play a functional role in directing attention: the processing of irrelevant information is attenuated by increasing alpha power in areas involved with processing this information (Foxe, Simpson, & Ahlfors, 1998; Hanslmayr et al., 2007; Jensen & Mazaheri, 2010).

L566-569: As such, it is conceivable that alpha oscillations can in some cases inhibit local transmission, while in other cases, depending on network location, connectivity and demand, alpha oscillation can facilitate signal transmission. This mechanism allows to increase transmission of relevant information and to block transmission of distractors.

(4) In the abstract and Figure 1, the authors claim an alternative function for alpha oscillations; that alpha "orchestrates signal transmission to later stages of the processing stream." In support, the authors cite their result showing that increased alpha activity originating from early visual cortex is related to enhanced visual processing in higher visual areas and association areas. This does not constitute a strong support for the alternative hypothesis. The correlation between posterior alpha power and frequency-tagged activity was not specific in any way; Fig. 10 shows that the correlation appeared on both 1) anticipating-auditory and anticipating-visual trials, 2) the visual tagged frequency and the auditory tagged activity, and 3) was not specific to the visual processing stream. Thus, the data is more parsimonious with a correlation than a causal relationship between posterior alpha and visual processing.

Again, the reviewer raises important points, which we want to address

The correlation between posterior alpha power and frequency-tagged activity was not specific, as it is present both when auditory and visual targets are expected:

If there is a connection between posterior alpha activity and higher-order visual information transfer, then it can be expected that this relationship remains across conditions and that a higher alpha activity is accompanied by higher frequency-tagged activity, both over trials and over conditions. However, it is possible that when alpha activity is lower, such as when expecting a visual target, the signal-to-noise ratio is affected, which may lead to higher difficulty to find a correlation effect in the data when using non-invasive measurements.

The connection between alpha activity and frequency-tagged activity appears both for auditory as well as visual stimuli and The correlation is not specific to the visual processing stream:

While we do see differences between conditions (e.g. in the EEG-analysis, mostly 36 Hz correlated with alpha activity and only in one condition 40 Hz showed a correlation as well), it is true that in our MEG analysis, we found correlations both between alpha activity and 36 Hz as well as alpha activity and 40 Hz.

We acknowledge that when analysing frequency-tagged activity on a trial-by-trial basis, where removal of non-timelocked activity through averaging (which we did when we tested for condition differences in Fig. 4 and 9) is not possible, there is uncertainty in the data. Baseline-correction can alleviate this issue, but it cannot offset the possibility of non-specific effects. We therefore decided to repeat the analysis with a fast-fourier calculated power instead of the Hilbert power, in favour of a higher and stricter frequency-resolution, as we averaged over a time-period and thus, the time-domain was not relevant for this analysis. In this more conservative analysis, we can see that only 36 Hz tagged activity when expecting an auditory target correlated with early visual alpha activity.

Additionally, we added correlation analyses between alpha activity and frequency-tagged activity within early visual areas, using the sensor cluster which showed significant condition differences in alpha activity. Here, no correlations between frequency-tagged activity and alpha activity could be found (apart from a small correlation with 40 Hz which could not be confirmed by a median split; see SUPPL Fig. 14 C). The absence of a significant correlation between early visual alpha and frequency-tagged activity has previously been described by others (Zhigalov & Jensen, 2020) and a Bayes factor of below 1 also indicated that the alternative hypotheses is unlikely.

Nonetheless, a correlation with auditory signal is possible and could be explained in different ways. For example, it could be that very early auditory feedback in early visual cortex (see for example Brang et al., 2022) is transmitted alongside visual information to higher-order areas. Several studies have shown that alpha activity and visual as well as auditory processing are closely linked together (Bauer et al., 2020; Popov et al., 2023). Inference on whether or how this link could play out in the case of this manuscript expands beyond the scope of this study.

To summarize, we believe the fact that 36 Hz activity within early visual areas does not correlate with alpha activity on a trial-by-trial basis, but that 36 Hz activity in other areas does, provides strong evidence that alpha activity affects down-stream signal processing.

We mention this analysis now in our discussion:

L533-536: Our data provides evidence in favour of this view, as we can show that early sensory alpha activity does not covary over trials with SSEP magnitude in early visual areas, but covaries instead over trials with SSEP magnitude in higher order sensory areas (see also SUPPL. Fig. 14).

Reviewer #1 (Recommendations for the authors):

The evidence for the alternative hypothesis, that alpha in early sensory areas orchestrates downstream signal transmission, is not strong enough to be described up front in the abstract and Figure 1. I would leave it in the Discussion section, but advise against mentioning it in the abstract and Figure 1.

We appreciate the reviewer’s concern regarding the inclusion of the alternative hypothesis—that alpha activity in early sensory areas orchestrates downstream signal transmission—in the abstract and Figure 1. While we agree that this interpretation is still developing, recent studies (Keitel et al., 2025; Clausner et al., 2024; Yang et al., 2024) provide growing support for this framework.

In response, we have revised the introduction, discussion, and Figure 1 to clarify that our intention is not to outright dismiss the alpha inhibition hypothesis, but to refine and expand it in light of new data. This revision does not invalidate the prior literature on alpha timing and inhibition; rather, it proposes an updated mechanism that may better account for observed effects.

We have though retained Figure 1, as it visually contextualizes the broader theoretical landscape. while at the same time added further analyses to strengthen our empirical support for this emerging view.

References:

Bastos, A. M., Litvak, V., Moran, R., Bosman, C. A., Fries, P., & Friston, K. J. (2015). A DCM study of spectral asymmetries in feedforward and feedback connections between visual areas V1 and V4 in the monkey. NeuroImage, 108, 460–475. https://doi.org/10.1016/j.neuroimage.2014.12.081

Bauer, A. R., Debener, S., & Nobre, A. C. (2020). Synchronisation of Neural Oscillations and Cross-modal Influences. Trends in cognitive sciences, 24(6), 481–495. https://doi.org/10.1016/j.tics.2020.03.003

Brang, D., Plass, J., Sherman, A., Stacey, W. C., Wasade, V. S., Grabowecky, M., Ahn, E., Towle, V. L., Tao, J. X., Wu, S., Issa, N. P., & Suzuki, S. (2022). Visual cortex responds to sound onset and offset during passive listening. Journal of neurophysiology, 127(6), 1547–1563. https://doi.org/10.1152/jn.00164.2021

Clausner T., Marques J., Scheeringa R. & Bonnefond M (2024). Feature specific neuronal oscillations in cortical layers BioRxiv :2024.07.31.605816. https://doi.org/10.1101/2024.07.31.605816

Diederich, A., & Colonius, H. (2008). When a high-intensity "distractor" is better then a low-intensity one: modeling the effect of an auditory or tactile nontarget stimulus on visual saccadic reaction time. Brain research, 1242, 219–230. https://doi.org/10.1016/j.brainres.2008.05.081

Haegens, S., Nácher, V., Luna, R., Romo, R., & Jensen, O. (2011). α-Oscillations in the monkey sensorimotor network influence discrimination performance by rhythmical inhibition of neuronal spiking. Proceedings of the National Academy of Sciences of the United States of America, 108(48), 19377–19382. https://doi.org/10.1073/pnas.1117190108

Jacoby, O., Hall, S. E., & Mattingley, J. B. (2012). A crossmodal crossover: opposite effects of visual and auditory perceptual load on steady-state evoked potentials to irrelevant visual stimuli. NeuroImage, 61(4), 1050–1058. https://doi.org/10.1016/j.neuroimage.2012.03.040

Keitel, A., Keitel, C., Alavash, M., Bakardjian, K., Benwell, C. S. Y., Bouton, S., Busch, N. A., Criscuolo, A., Doelling, K. B., Dugue, L., Grabot, L., Gross, J., Hanslmayr, S., Klatt, L.-I., Kluger, D. S., Learmonth, G., London, R. E., Lubinus, C., Martin, A. E., … Kotz, S. A. (2025). Brain rhythms in cognition – controversies and future directions. ArXiv. https://doi.org/10.48550/arXiv.2507.15639

Nickerson R. S. (1973). Intersensory facilitation of reaction time: energy summation or preparation enhancement?. Psychological review, 80(6), 489–509. https://doi.org/10.1037/h0035437

Popov, T., Gips, B., Weisz, N., & Jensen, O. (2023). Brain areas associated with visual spatial attention display topographic organization during auditory spatial attention. Cerebral cortex (New York, N.Y. : 1991), 33(7), 3478–3489. https://doi.org/10.1093/cercor/bhac285

Salagovic, C. A., & Leonard, C. J. (2021). A nonspatial sound modulates processing of visual distractors in a flanker task. Attention, perception & psychophysics, 83(2), 800–809. https://doi.org/10.3758/s13414-020-02161-5

Scheeringa, R., Petersson, K. M., Kleinschmidt, A., Jensen, O., & Bastiaansen, M. C. (2012). EEG α power modulation of fMRI resting-state connectivity. Brain connectivity, 2(5), 254–264. https://doi.org/10.1089/brain.2012.0088

Spaak, E., Bonnefond, M., Maier, A., Leopold, D. A., & Jensen, O. (2012). Layer-specific entrainment of γ-band neural activity by the α rhythm in monkey visual cortex. Current biology : CB, 22(24), 2313–2318. https://doi.org/10.1016/j.cub.2012.10.020

Yang, X., Fiebelkorn, I. C., Jensen, O., Knight, R. T., & Kastner, S. (2024). Differential neural mechanisms underlie cortical gating of visual spatial attention mediated by alpha-band oscillations. Proceedings of the National Academy of Sciences of the United States of America, 121(45), e2313304121. https://doi.org/10.1073/pnas.2313304121

Zhigalov, A., & Jensen, O. (2020). Alpha oscillations do not implement gain control in early visual cortex but rather gating in parieto-occipital regions. Human brain mapping, 41(18), 5176–5186. https://doi.org/10.1002/hbm.25183

Zumer, J. M., Scheeringa, R., Schoffelen, J. M., Norris, D. G., & Jensen, O. (2014). Occipital alpha activity during stimulus processing gates the information flow to object-selective cortex. PLoS biology, 12(10), e1001965. https://doi.org/10.1371/journal.pbio.1001965

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

Selective attention involves prioritizing relevant sensory input while suppressing irrelevant stimuli. It has been proposed that oscillatory alpha-band activity (∼10 Hz) aids this process by functionally inhibiting early sensory regions. However, recent studies have challenged this notion. Our EEG and MEG studies aimed to investigate whether alpha oscillations serve as a ‘gatekeeper’ for downstream signal transmission. We first observed these effects in an EEG study and then replicated them using MEG, which allowed us to localize the sources.

We employed a cross-modal paradigm where visual cues indicated whether upcoming targets required visual or auditory discrimination. To assess inhibition, we utilized frequency-tagging, simultaneously flickering the fixation cross at 36 Hz and playing amplitude-modulated white noise at 40 Hz during the cue-to-target interval.

Consistent with prior research, we observed an increase in posterior alpha activity following cues signalling auditory targets. However, remarkably, both visual and auditory frequency tagged responses amplified in anticipation of auditory targets, correlating with alpha activity amplitude. Our findings suggest that when attention shifts to auditory processing, the visual stream remains responsive and is not hindered by occipital alpha activity. This implies that alpha modulation does not solely regulate ‘gain control’ in early visual areas but rather orchestrates signal transmission to later stages of the processing stream.

Introduction

Earlier version: In our daily life, we are often confronted with sensory information from many different sources, all at once. To operate effectively, we require selective attention, reconciling the tension between environmental inputs relevant for top-down goals and sensory information that may be perceptually salient but task-irrelevant. Previous research gave rise to the prominent alpha inhibition hypothesis, which suggests that oscillatory activity in the alpha range (∼10 Hz) plays a mechanistic role in selective attention through functional inhibition of irrelevant cortices (see Fig. 1A; Foxe et al., 1998; Jensen & Mazaheri, 2010; Klimesch et al., 2007). Functional inhibition refers to an area of the cortex being actively hindered to process input, which is distinctly different from idling, where a part of the cortex is not actively involved. Evidence supporting this theory revealed an increase in alpha-power over task-irrelevant sensory cortices after the onset of cues indicating the spatial location (Kelly et al., 2006; Okazaki et al., 2014; Thut et al., 2006; Worden et al., 2000; Zumer et al., 2014) or specific modality of an upcoming target (Foxe et al., 1998; Fu et al., 2001; Mazaheri et al., 2014). Moreover, previous investigations have observed ‘spontaneous fluctuations’ of alpha power in sensory regions, particularly in the visual cortex, to be inversely related to discrimination ability (Ergenoglu et al., 2004; Van Dijk et al., 2008). Alpha inhibition is believed to be transmitted in a phasic manner, as phosphene perception as well as high-frequency and spiking activity vary in line with the alpha cycle (Dugué et al., 2011; Haegens et al., 2011; Spaak et al., 2012).

Revised version: In our daily life, we are often confronted with sensory information from many different sources, all at once. To operate effectively, we require selective attention, reconciling the tension between environmental inputs relevant for top-down goals and sensory information that may be perceptually salient but task-irrelevant. Previous research gave rise to the prominent alpha inhibition hypothesis, which suggests that oscillatory activity in the alpha range (∼10 Hz) plays a mechanistic role in selective attention through functional inhibition of irrelevant cortical areas (see Fig. 1; Foxe et al., 1998; Jensen & Mazaheri, 2010; Klimesch et al., 2007). Functional inhibition refers to an area of the cortex being actively hindered to process input, which is distinctly different from idling, where a part of the cortex is not actively involved. Evidence supporting this theory revealed an increase in alpha-power over task-irrelevant sensory cortices after the onset of cues indicating the spatial location (Kelly et al., 2006; Okazaki et al., 2014; Thut et al., 2006; Worden et al., 2000; Zumer et al., 2014) or specific modality of an upcoming target (Foxe et al., 1998; Fu et al., 2001; Mazaheri et al., 2014). Moreover, previous investigations have observed ‘spontaneous fluctuations’ of alpha power in sensory regions, particularly in the visual cortex, to be inversely related to discrimination ability (Ergenoglu et al., 2004; Van Dijk et al., 2008). Alpha inhibition is believed to be transmitted in a phasic manner, as phosphene perception as well as high-frequency and spiking activity vary in line with the alpha cycle (Dugué et al., 2011; Haegens et al., 2011; Spaak et al., 2012).

Illustration of the alpha inhibition hypothesis.

Earlier version: The alpha inhibition hypothesis suggests that occipital alpha inhibits visual information processing in a phasic manner. If alpha activity is high, it suggests that the whole area is inhibited and thereby disengaged (Foxe et al., 1998; Jensen & Mazaheri, 2010; Klimesch et al., 2007). We propose a revision of this hypothesis, whereby alpha activity exerts its phasic inhibition to regulate downstream information transfer, creating enhanced signal packages of prioritised information (see also Yang et al., 2023; Zhigalov & Jensen, 2020; Zumer et al., 2014).

Revised version: The alpha inhibition hypothesis suggests that occipital alpha inhibits visual information processing in a phasic manner. If alpha activity is high, it suggests that neural processing in the cortical area is inhibited (Foxe et al., 1998; Jensen & Mazaheri, 2010; Klimesch et al., 2007). We propose a revision of this hypothesis, whereby alpha activity exerts its phasic inhibition to regulate information transfer, creating enhanced signal packages of prioritised information (see also Yang et al., 2023; Zhigalov & Jensen, 2020; Zumer et al., 2014). This way, irrelevant or distracting information is inhibited through a block of transfer, rather than through blocking of incoming sensory information.

Recent evidence challenged a direct connection between alpha activity and visual information processing in early visual cortex. As such, both visual steady-state responses and alpha power were modulated by attention but did not covary when investigating individual trials (Zhigalov & Jensen, 2020). Unfortunately, very few studies have investigated direct connections between alpha activity, attention and sensory signals, especially over trials. Furthermore, results seem to depend on timing of alpha activity in relation to sensory responses as well as stimulus type and outcome measure (Morrow et al., 2023).

Earlier version: Accordingly, the objective of the current study is to test the alpha inhibition hypothesis compared to an alternative theory. Based on the alpha inhibition hypothesis, alpha modulation is connected to ‘gain control’ in early visual areas through modulation of excitability (Foxe & Snyder, 2011; Jensen & Mazaheri, 2010; Van Diepen et al., 2019). In contrast, we propose that inhibitory effects of alpha modulation are exhibited at later stages of the processing stream (Peylo et al., 2021; Yang et al., 2023; Zhigalov & Jensen, 2020; Zumer et al., 2014), gating feedforward or feedback communication between sensory areas (see Fig. 1; Bauer et al., 2020; Haegens et al., 2015; Uemura et al., 2021).

Revised version: Accordingly, the objective of the current study is to test the alpha inhibition hypothesis compared to an alternative theory. Based on the alpha inhibition hypothesis, alpha modulation is connected to ‘gain control’ in early visual areas through modulation of excitability (Foxe & Snyder, 2011; Jensen & Mazaheri, 2010; Van Diepen et al., 2019). In contrast, we propose that functional and inhibitory effects of alpha modulation, such as distractor inhibition, are exhibited through blocking or facilitating signal transmission to higher order areas (Peylo et al., 2021; Yang et al., 2023; Zhigalov & Jensen, 2020; Zumer et al., 2014), gating feedforward or feedback communication between sensory areas (see Fig. 1; Bauer et al., 2020; Haegens et al., 2015; Uemura et al., 2021).

To this end, we applied frequency-tagging, the rhythmic presentation of sensory stimuli, which elicits steady-state sensory evoked potentials or fields (SSEP/SSEF), consisting of rhythmic neuronal activity in the frequency of stimulation (Brickwedde et al., 2020; Colon et al., 2012; Dinse et al., 1998; Marzoll et al., 2018; Regan, 1982; Snyder, 1992; Stapells et al., 1984; Tobimatsu et al., 1999). The magnitude of SSEPs is attention-dependent (de Jong et al., 2010; Müller et al., 1998; Müller & Hillyard, 2000; Porcu et al., 2013; Saupe et al., 2009; Toffanin et al., 2009) even for frequencies too fast to perceive consciously (Brickwedde et al., 2022; Zhigalov et al., 2019). Visual SSEP/SSEFs are most strongly observable over occipital areas, whereas auditory SSEP/SSEFs appear most strongly over temporal (MEG) or fronto-to-central (EEG) areas (de Jong et al., 2010; Hari et al., 1989; Pantev et al., 1996; Regan, 1982). Furthermore, when applying two different frequencies for two different sensory modalities, their intermodulation frequency (f1-f2) has been suggested to reflect cross-modal integration (Drijvers et al., 2021). Due to distinct responses, localisation and attention-dependence, frequency-tagging provides an optimal tool to study sensory signal processing and integration over time.

The aim of our study was to directly test the alpha inhibition hypothesis by investigating if cue-induced modulation of alpha activity coincides with the suppression of frequency-tagging responses in task-irrelevant modalities. Based on previous studies, we utilized a cross-modal attention paradigm, in which symbolic visual cues signalled the target modality (visual or auditory) of an upcoming discrimination task (e.g., van Diepen et al., 2015). Here, we included an additional experimental manipulation in the form of frequency-tagging to assess the involvement of the auditory and visual systems in the cue-to-target interval.

In line with previous results, we hypothesized that signalling an upcoming auditory target would lead to increased alpha activity over visual occipital regions as well as increased SSEP responses to auditory and decreased SSEP to visual stimuli as indexed by frequency-tagging. In brief, while we observed the expected cue-induced early-visual alpha modulation, the amplitude of auditory and visual SSEP/SSEFs as well as their intermodulation frequency increased just prior to the onset of the auditory target, contradicting the alpha inhibition hypothesis. The difference between conditions of visual SSEP/SSEFs originated from sensory integration areas and correlated with early sensory alpha activity on a trial-by-trial basis, speaking to an effect of alpha modulation on signal transmission rather than inhibition of early visual areas.

Removed in the revised version: Results

Results Study 1 – cross-modal EEG experiment

To assess audio-visual excitability in anticipation of either visual or auditory targets, we displayed visual cues to signal the modality of the upcoming target (auditory, visual, or non-specific). In a three-second cue-to-target interval, we frequency-tagged the fixation cross at 36 Hz and played 40 Hz amplitude modulating white noise. Immediately afterwards, participants had to either discriminate between three different pitch sounds (auditory target) or three different Gabor patch orientations (visual target). If the target modality was cued (e.g. not non-specific), 50% of the trials were accompanied by a random distractor from the target pool of the opposing sensory modality (see Fig. 2A).

A, Illustration of the cross-model discrimination task in Study 1. Trials were separated by a 4 s interval, in which a fixation cross was displayed. A brief central presentation of the cue (100 ms) initiated the trial, signalling the target modality (see figure above from left to right: auditory, non-specific, visual). In the cue-to-target interval, the fixation cross was frequency-tagged at 36 Hz. At the same time, a sound was displayed over headphones, which was frequency-tagged at 40 Hz. Both tones and fixation cross contained no task-relevant information. The target, consisting either of a static Gabor patch or a tone, was presented for 25 ms. Participants had to differentiate via button presses between 3 different Gabor rotations or tone pitches, respectively. In 50% of auditory and visually cued trials, a distractor in form of a random pitch or rotation of the un-cued modality was presented alongside the target. B-C, Analysis of task accuracy and reaction time indicates increased difficulty of auditory targets. B, Task accuracy compared between all 6 experimental conditions reveals a drop in accuracy for responses to auditory targets. C, reaction times of correct trials compared between all 6 experimental conditions. The slowest reaction times are observable following auditory targets alongside visual distractors. N = 22; *** sig < .001; ** sig. < .01; * sig. < .05;

Behavioural performance

We found that accuracy differed significantly between conditions (F(5,105) = 44.16; p < .001). Overall, participants were significantly less accurate in the auditory discrimination task (‘overall auditory’, M = 79% correct, SD = 11.7) than in the visual discrimination task (‘overall visual’, M = 97%, SD = 2.3; see Fig. 2B), with the worst performance occurring when auditory targets were paired with visual distractors (auditory +: M = 74% correct, SD = 13).

Reaction times yielded a similar pattern, with auditory reactions (overall auditory: M = 662 ms, SD = 136) being slower than for the visual task (overall visual: M = 597 ms, SD = 130; main effect over all conditions: F(5,105) = 27.47; p < .001; see Fig. 2C). This was mostly driven by slow responses to auditory targets paired with distractors (auditory +: M = 723 ms).

To ascertain that cues and distractors were functionally relevant, we calculated attentional benefit (cued - non-specific condition) and distractor cost (cued condition with distractor – cued condition without distractor). As expected, cues improved reaction times over all conditions and distractors impaired the accuracy over all conditions and reaction times for auditory targets. (see suppl Fig. 1). Interestingly, auditory distractors reduced reaction times to visual targets, which could be explained by a generally faster processing of auditory targets (Jain et al., 2015). As such, the auditory distractor possibly caused intersensory facilitation (Nickerson, 1973), whereby reaction times to a target can be facilitated when accompanied by stimuli of other sensory modalities, even if they are irrelevant or distracting.

Cross-modal cues differentially modulated pre-target alpha activity

We conducted a time-frequency analysis of power in the cue-to-target interval and found a stronger amplitude increase from baseline for auditory compared to visual target conditions starting around 2 s before target onset (Fig. 3). We calculated the time course of alpha power changes using the Hilbert-transformation (8 – 12 Hz, Fig. 3A-B). We then applied cluster permutation analysis, whereby real condition differences were tested against coincidental findings by randomly permutating the condition labels to the data and testing for condition differences 1000 times (Maris & Oostenveld, 2007). Consistent with previous work (Mazaheri et al., 2014; van Diepen & Mazaheri, 2017), analysing the last two seconds before target onset revealed two significant clusters of difference in alpha power when expecting an auditory compared to a visual target. These effects corresponded to clusters extending from −1.84 to – 0.64 s (p = .004) and −0.62 to 0s (p = .005). When presenting non-specific cues, alpha power changes were not significant, but descriptively larger compared to visual target conditions and lower compared to auditory target conditions (see suppl Fig. 2). However as significant alpha modulation was a prerequisite to test our hypotheses, we excluded this condition from further analysis.

Post-cue modality specific modulation of alpha power in anticipation of an auditory versus a visual target

A-B, the time course of post-cue alpha power. Cluster permutation analysis resulted in two condition effects, both indicating heightened alpha activity when expecting an auditory compared to a visual target (A: p < .01; B: p < .01). C-D, Time-frequency representation of power in the cue-to-target interval. A greater increase in alpha power was observed when expecting an auditory target (average over significant electrodes for the condition difference in A). note: Cluster electrodes are marked in white. Shading represents standard error from the mean; Δ / ∑ represents (a-b)/(a+b) normalization.

Cross-modal cues increased the amplitude of the frequency-tagged responses across both modalities

To assess the temporal development of frequency-tagging responses, steady-state potentials were calculated using data band-pass filtered around the tagging frequency. Neuronal 40 Hz auditory-steady state evoked potential (ASSEP) responses were strongest over central areas (Fig. 4A), and 36 Hz auditory-steady state evoked potential (VSSEP) responses were strongest over occipital areas (Fig. 4B).

Increase in amplitude of both visual and auditory frequency tagged responses when anticipating visual or auditory targets.

Event-related potentials and scalp topographies reveal distinct modality specific responses at the tagged frequencies. A, auditory steady-state evoked potential (ASSEP) averaged over 6 central electrodes displaying the highest 40 Hz power (Cz, FC1, FC2, C1, C2, FCz; marked with white squares). B, visual steady-state evoked potential (VSSEP) averaged over 4 occipital electrodes displaying the highest 36 Hz power (POz, O1, O2, Oz; marked with white squares) C, the Hilbert-envelope of the 40 Hz ASSEP reveals an increase shortly before target onset when anticipating an auditory compared to a visual target (p = .041); D, the Hilbert-envelope of the 36 Hz VSSEP likewise reveals an increase shortly before target onset when anticipating an auditory compared to a visual target (p = .014). note: Cluster electrodes are marked in white. Shading represents standard error from the mean. Δ / ∑ represents (a-b)/(a+b) normalization.

To assess the differences between conditions, the Hilbert envelope of the steady-state potentials was analysed using cluster permutation analyses. When expecting an auditory target compared to a visual target, the 40 Hz ASSEP response was larger shortly before target onset (see Fig. 4C; −0.15 to −0.08 s; p = .041).

Surprisingly, the 36 Hz VSSEP response was likewise increased shortly before expecting an auditory compared to a visual target (see Fig. 4D; −0.16 to −0.06 s; p = .014). As such, the observed frequency tagging responses might reflect effort, which affects the vigilance of the whole sensory system rather than sensory-specific allocations of attention. For both 36 Hz VSSEP and 40 Hz ASSEP responses, condition differences appeared strongest over mid-parietal regions in contrast to primary sensory cortices.

Alpha power was positively correlated with amplitude of frequency tagged responses

Following the observation of condition differences in alpha activity and frequency-tagging responses, we were further interested in exploring whether these responses were connected. Accordingly, we conducted trial-by-trial correlations using alpha condition differences and their electrode positions as seed, which was correlated with frequency-tagging signals over all electrodes. Multiple comparison correction was applied by testing the correlation matrix against a zero-correlation matrix with a cluster permutation approach. A positive correlation was observed over right parietal-to-occipital areas between the late alpha cluster activity and both 40 Hz ASSEP (p = .009) and 36 Hz VSSEP (p = .004) responses when expecting a visual target (see Fig. 5.A-B). To estimate the robustness of these results, we additionally conducted median split analyses between trials with high and low alpha power for each participant, as well as averaged the correlation coefficient of each participant and calculated a one-sample t-test against 0. For each analysis we provided the Bayes Factor, which estimates the strength of support for or against the null hypothesis (BF > 3.2 is considered as substantial evidence and BF > 10 is considered as strong evidence; Kass & Raftery, 1995).

Relationship between cue induced alpha modulation and amplitude of frequency tagged responses in study 1.

Previously obtained alpha clusters (see Fig. 3) were correlated over trials with auditory 40 Hz and visual 36 Hz clusters (see Fig. 4), where alpha electrodes/sensors were applied as seeds. The analysis was performed using a cluster-permutation approach, testing a correlation model against a 0-correlation model. Clusters significantly diverging from the 0-correlation model are presented topographically. Additionally, median splits between high and low alpha trials as well as correlation coefficients of these clusters are displayed for all participants A-B, a positive correlation is visible between alpha activity in the last 400 ms and steady state potentials shortly before target onset when expecting a visual target (visual 36 HZ: p = .013; 40 Hz: p = .009). D, when expecting an auditory target, there is a positive correlation between alpha activity in the last 400 ms and visual 36 Hz activity shortly before target onset (p = .010). E, the correlation between alpha activity 400 ms and visual 36 Hz activity shortly before target onset changes its direction depending on whether an auditory or a visual target is expected (p = .037). C, a positive correlation is also visible between alpha activity as early as ∼1200 ms to 400 ms and visual 36 Hz activity shortly before target onset when expecting a visual target (p = .016). N = 22; *** sig < .001; ** sig < .01; * sig < .05. + sig. < .1

The median split was highly significant for the 36 Hz VSSEP response (Fig. 5A, middle columns, p = .033; t(19) = 2.29; BF(10) = 1.91) but did not reach significance for the 40 Hz ASSEP response (Fig. 5B, middle column; p = 0.20; t(19) = 1.32; BF(10) = 0.49). Correlation coefficients indicated strong correlations with alpha activity for both 40 Hz ASSEP (Fig. 5B, right column; p < .001; t(19) = 4.95; BF(10) = 306.93) and 36 Hz VSSEP activity (Fig. 5A, right column; p < .01; t(19) = 3.66; BF(10) = 23.57).

The same positive correlation with alpha activity was found when expecting an auditory target for 36 Hz VSSEP activity (p = .031), but not for 40 Hz ASSEP activity (see Fig. 5D). A median split between high and low alpha activity (p = .005; t(19) = 3.14; BF(10) = 8.53) and correlation coefficients (p = .002; t(19) = 3.52; BF(10) = 17.76) provided moderate to strong evidence for this effect. The likely origin of alpha modulations from early visual cortices may explain a stronger connection to 36 Hz VSSPE responses compared to 40 Hz ASSEP responses.

Additionally, we compared how correlation coefficients between alpha activity and frequency-tagging differed when anticipating an auditory versus a visual target. Multiple comparison correction was applied with cluster permutation analysis. Interestingly, an interaction between the strength of the correlation associating alpha, 36 Hz VSSEP activity and condition became apparent (p = .044; see Fig. 5E) and was observed most strongly over right-central electrodes. Comparing the correlation coefficients of participants over this cluster revealed a strong effect and even a change of direction in the correlation (p < .001; t(21) = −4.76; BF(10) = 259.97). Particularly, when expecting a visual target, there was a negative correlation between 36 Hz VSSEP and alpha activity, which turned positive when expecting an auditory target. Both correlations also differed significantly from 0 (expecting a visual target: p < .001; t(21) = −3.87; BF(10) =39.62; expecting an auditory target: p = .020; t(21) = 2.53; BF(10) =2.83). In contrast to the previously observed positive correlation between alpha activity and 36 Hz VSSEP activity, the significant electrode cluster was located more ventrally. This effect could possibly hint at dynamic adaptability of oscillatory alpha effects on later processing stages.

It is further noteworthy that the correlation between alpha activity and 36 Hz VSSEP response when expecting a visual target, was also present when utilizing the early alpha cluster as seed (∼ 1200 to 400 ms before target onset, see Fig. 5C), in which case alpha modulation greatly preceded the 36 Hz VSSEP modulation (p = .016). For this correlation, the median split between high and low alpha trials did not reach significance (p = .11; t(20) = 1.69; BF(10) =0.76), however testing the correlation coefficients against 0 again revealed a significant effect (p = .003; t(20) = 3.39; BF(10) =14.22). This result may support the directionality of alpha modulation affecting visual information processing (and not vice versa).

Intermodulation frequency

Earlier version: Lastly, we analysed the steady-state response of the intermodulation frequency at 4 Hz. Increased intermodulation shortly before target onset could be observed when expecting an auditory compared to a visual target (−0.51 to −0.0620 s; p < .001). This effect was strongest over left fronto-to-central electrodes and right central-to-occipital electrodes (see Fig. 6A) and may suggest in increase in cross-modal integration.

Revised version: Lastly, we analysed the steady-state response of the intermodulation frequency at 4 Hz. Increased intermodulation shortly before target onset could be observed when expecting an auditory compared to a visual target (−0.51 to −0.0620; p < .001). This effect was strongest over left fronto-to-central electrodes and right central-to-occipital electrodes (see Fig. 6A) and may suggest in increase in cross-modal integration.

Steady-state response in the intermodulation frequency and its behavioural relevance.

A, the Hilbert-envelope of the 4 Hz steady-state response reveals an increase shortly before target onset when anticipating an auditory compared to a visual target (p < .01). B, there is a trial-by-trial correlation between 4 Hz activity and reaction time when a visual target without distractor was presented. The correlation is further illustrated by a median split between fast and slow reaction time trials as well as by correlation coefficients for each participant. N = 22; ** sig. < .01;

With the goal to examine whether there are any behavioral effects of the condition difference, trial by trial correlations with reaction times of each of the 2 auditory and visual conditions were performed. Only in the easiest condition, when expecting a visual target that was not accompanied by a distractor, a negative correlation with reaction time could be found, strongest over right central electrodes (see Fig. 6B; p = .046). While a median split between slow and fast trials did not reach significance (p = .50; t(21) = −0.69; BF(10) = .28), testing the correlation coefficients against 0 revealed strong evidence for a correlation (p = .004; t(21) = −3.28; BF(10) = 11.98). If the intermodulation frequency reflects cross-modal integration (Drijvers et al., 2021), this effect would indicate faster reaction times for trials with stronger sensory integration.

Results Study 2 – cross-modal MEG experiment

To test the robustness of our results and to employ additional control analyses, we replicated our experiment using MEG (see Fig. 7A). While an increase in visual information processing parallel to an increase in alpha modulation already contradicts the notion of alpha inhibition exerting “gain control”, affecting the whole visual cortex, our claim that alpha modulation instead affects visual information at later processing stages still required further validation. As such, our goal was to perform source analyses showing alpha modulation originating from primary visual areas affected visual information at later processing stages (e.g. not in primary visual cortex). Additionally, to exclude that the uncertainty over possible distractors affected our results, we employed a block design, where block 1 consisted only of trials without distractors and in block 2 targets were always accompanied by a distractor. Furthermore, we aligned the visual and auditory task to be more similar, both of them now featuring frequency-discrimination, which related to sound pitch (frequency) in the auditory condition and stripe-frequency of the Gabor patch in the visual condition. Lastly, to make sure our effects were driven by sensory modality-differences rather than task-difficulty differences, we included a short calibration phase. Prior to the experiment, difficulty of pitch sounds, and Gabor patch frequency were calibrated for each individual, ascertaining a success rate between 55% to 75%.

A, Illustration of the cross-model discrimination task in Study 2. Trials were separated by a 1-1.3 s interval, in which a fixation cross was displayed. A brief central presentation of the cue (100 ms) initiated the trial, signalizing the target modality (from left to right: auditory, visual). In the cue-to-target interval, the fixation cross was frequency-tagged at 36 Hz. At the same time, a sound was displayed over headphones, which was frequency-tagged at 40 Hz. Both tones and fixation cross contained no task-relevant information. The target, consisting either of a static Gabor patch or a tone, was presented for 25 ms. Participants had to differentiate via button presses between 3 different Gabor stripe frequencies or tone pitches, respectively. In block 2, a distractor in form of a random sound or Gabor patch of the un-cued modality was always presented alongside the target. B-C, analysis of task accuracy and reaction time indicates comparable difficulties in block 2. B, task accuracy differences were only observable in the first block without distractors. C, reaction times to visual targets in the first block were strongly increased compared to all other conditions. In the second block, no significant difference in reaction times was observable. N = 27; *** sig < .001; ** sig. < .01; * sig. < .05;

Behavioural performance

Our adjustments in the MEG-study streamlined performances to be more in line between auditory and visual conditions in the second block, however in the first block there were strong condition differences (Main effect condition: (F(3,75) = 10.26; p < .01). In block 1, participants were significantly more accurate in the auditory (‘block 1 auditory’, M = 84% correct, SD = 10.03) compared to the visual task (‘block 1 visual’, M = 70%, SD = 10.61; see Fig. 7B). However, in block 2, there was no observable difference between visual and auditory task accuracy (‘block 2 auditory’, M = 77% correct, SD = 10.03; (‘block 2 visual’, M = 78%, SD = 13.15). Reaction times were comparable between conditions (overall: M = 661 ms, SD = 76), with the exception of responses to visual targets in block 1 (block 1 visual: M = 726 ms, SD = 85), which were significantly slower than all other reaction times (main effect condition: F(3,75) = 18.13; p < .001; see Fig. 7C).

Due to the block design, distractor cost could not be differentiated from learning effects. For auditory targets, there was a significant decrease in accuracy between block 1 and 2, while for visual targets, there was a significant increase. Similar to study 1, reaction times for auditory targets stayed roughly the same, but became shorter for visual targets (see suppl Fig. 1).

Cross-modal cues differentially modulated pre-target alpha activity in early visual areas

In line with previous studies (van Diepen & Mazaheri, 2017), condition differences in alpha activity were only significant in block 2, where distractors were always present. As alpha modulation was a prerequisite to test our hypotheses, we performed the following analyses solely with data from block 2 (see Fig. 8).

Post-cue modality specific early visual modulation of alpha power in anticipation of an auditory versus a visual target

A, the time course of post-cue alpha power. Cluster permutation analysis resulted in a condition effects, indicating heightened alpha activity when expecting an auditory compared to a visual target (p = .034). B, source localization of the condition difference between expecting an auditory versus a visual target, revealing a significant cluster in early visual areas with stronger effects on the right hemisphere (p < .01). C-D, Time-frequency representation of power in the cue-to-target interval (average over electrodes that showed maximal condition difference in A). note: Cluster electrodes are marked in white. Shading represents standard error from the mean; Δ / ∑ represents (a-b)/(a+b) normalization.

Replicating our results from study 1, cluster permutation on the Hilbert-transformed alpha activity revealed a significant cluster between −1.47 to −1.18 s (p = .034), where alpha change from baseline was higher when expecting an auditory compared to a visual target (see Fig. 8A). Applying roughly the cluster time-window (−1.5 and −1 s), we conducted a source localization, contrasting the two conditions with cluster permutation analysis. We found that condition differences in alpha activity originated from early visual areas, albeit with a stronger effect on the right compared to the left hemisphere (p < .01; peak spm coordinates: 41 −82 −19 mm, in the right lingual gyrus see Fig. 8B). Additionally, we replicated this effect with a virtual channel analysis in V1 (see SUPPL Fig. 12)

Cross-modal cues increased the amplitude of the frequency-tagged responses across both modalities

In accordance with study 1, neuronal responses to 40 Hz auditory steady-state evoked field (ASSEF) responses were strongest over temporal areas (Fig. 9A) and 36 Hz visual steady-state evoked field (VSSEF) responses were strongest over occipital areas (Fig. 9C). As expected, the auditory tagging response originated from the right-hemispheric early auditory cortex (cluster significance: p < .01; peak spm coordinates: 69 −22 5 mm, in the right superior temporal gyrus; see Fig. 9B) and the visual tagging response originated from the early visual cortex (cluster significance: p < .001; peak spm coordinates: 19 −105-11 mm, in the right lingual gyrus; see Fig. 9D). Additionally, we observed significantly reduced 40 Hz activity in the left-hemispheric visual-to-central cortex (cluster significance: p < .01).

Post-cue modality specific early visual modulation of alpha power in anticipation of an auditory versus a visual target.

A, auditory steady-state evoked fields (ASSEF) averaged over 24 temporal sensors displaying the highest auditory 40 Hz power (12 right, 12 left; marked in white squares). B, ASSEF source localization revealed a significant positive cluster in the right-hemispheric early auditory cortex (p < .001). C, visual steady-state evoked fields (VSSEF) averaged over 10 occipital sensors, displaying the highest visual 36 Hz power (marked in white squares). D, VSSEF source localization revealed a significant positive cluster in the early visual cortex (p < .001). E, the Hilbert-envelope of the 40 Hz ASSEF reveals an increase shortly before target onset when anticipating an auditory compared to a visual target (p = .043); F, condition differences in the 40 Hz ASSEF response did not reach significance in sensor space. G, the Hilbert-envelope of the 36 Hz VSSEF likewise reveals an increase shortly before target onset when anticipating an auditory compared to a visual target (p = .019). H, condition differences in the 36 Hz VSSEF response were significant over several areas of the visual stream, including most strongly the medial occipital cortex, the calcarine fissure, and the precuneus (p = .047); note: Cluster electrodes are marked in white. Shading represents standard error from the mean. Δ / ∑ represents (a-b)/(a+b) normalization.

Applying cluster permutation on the Hilbert-transformed frequency-tagging data over the last 500 ms before target onset replicated the results of study 1, where both 36 Hz VSSEF as well as 40 Hz ASSEF signals were significantly enhanced when expecting an auditory compared to a visual target (see Fig. 9E-G; auditory target: p = .043; visual target: p = .019).

Source localization confirmed the condition difference in 36 Hz VSSEF activity to originate from later stages of the processing stream, encompassing a wide range of areas, most strongly the medial occipital cortex and the precuneus (cluster significance: p = .047; peak spm coordinates: 3 −66 44 mm, in the left and right precuneus; see Fig. 9H). In source space, the effect was not significant for auditory 40 Hz activity (p = .11; peak spm coordinates: −9 −39 0 mm, in the left precuneus; see Fig. 9F). Furthermore, a virtual channel analysis in V1 and Heschl’s gyrus confirmed that there were no condition differences in primary visual and auditory areas (see SUPPL Fig. 12).

Alpha power was positively correlated with amplitude of frequency tagged responses

The positive correlations between alpha power and frequency-tagging amplitude observed in study 1 could be replicated. Alpha activity 500 ms before target onset correlated both with 36 Hz VSSEF and 40 Hz ASSEF activity when expecting a visual target (see Fig. 10A-B, 36 Hz response: cluster significance: p < .01; median split: p < .001; t(24) = 4.33; BF(10) = 127; t-test: p < .001; t(24) = 5.33; BF(10) = 1272; 40 Hz response: cluster significance: p < .001; median split: p < .001; t(24) = 7.05; BF(10) = 59443; t-test: p < .001; t(24) = 6.75; BF(10) = 30515).

Relationship between cue induced alpha modulation and amplitude of frequency tagged responses in study 2.

A-C, a positive correlation is visible between alpha activity in the last 500 ms as well as alpha activity in the last 1500ms–1000 ms and steady state potentials shortly before target onset when expecting a visual target (36 HZ late: p = .013; 40 Hz late: p = .009; 40 Hz early: p = 002). D-F, when expecting an auditory target, there is a positive correlation between alpha activity in the last 500 ms as well as alpha activity in the last 1500ms–1000 ms and steady state potentials shortly before target onset (36 HZ late: p < .001; 40 Hz late: p = .005; 36Hz early: p = 011). N = 27; *** sig < .001; ** sig. < .01; * sig. < .05. + sig. < .1

The same correlation was present when expecting an auditory target for both 36 Hz VSSEF activity (see Fig. 10. D-E; cluster significance: p < .001; median split: p = .001; t(23) = 3.62; BF(10) = 25.51; t-test: p < .001; t(23) = 4.60; BF(10) = 216) and 40 Hz ASSEF activity (cluster significance: p = .005; median split: p < .001; t(25) = 3.75; BF(10) = 36.40; t-test: p < .001; t(25) = 4.06; BF(10) = 73.61).

In accordance with the results of study 1, we then tested whether alpha activity preceding frequency-tagging activity showed similar correlations. Our data revealed that during the last 1 to 1.5 s before target onset alpha activity correlated with the 36 Hz VSSEF response during the last 500 ms prior to an auditory target (see Fig. 10F; cluster significance: p = .01; median split: p < .001; t(24) = 4.66; BF(10) = 271; t-test: p < .001; t(24) = 5.66; BF(10) = 2688). The same alpha activity correlated with 40 Hz ASSEF activity during the last 500 ms prior to a visual target (see Fig. 10C; cluster significance: p = .002; median split: p = .002; t(23) = 3.41; BF(10) = 16.38; t-test: p < .001; t(23) = 5.57; BF(10) = 1901).

Lastly, both alpha activity as well as 36 Hz VSSEF responses 500 ms before target onset correlated negatively with reaction time on a trial-by-trial basis, indicating faster reaction times in trials with higher pre-stimulus activity (alpha: p = .037; median split: p = .013; t(25) = −2.67; BF(10) = 3.78; t-test: p < .01; t(25) = −3.34; BF(10) = 14.84; 36 Hz: p = .002; median split: p = .004; t(25) = −3.20; BF(10) = 8.98; t-test: p < .01; t(25) = −3.46; BF(10) = 19.12. See suppl Fig. 3-4).

The same increase in the intermodulatory frequency observed in study 1, could be observed in our second study, during the last 500 ms prior to target onset (see Fig. 11A; p = .006). In source space, a descriptive condition difference was visible in auditory and visual sensory cortices, however this effect did not reach significance (p = .49; peak spm coordinates: 45 −83 −19 mm, in the right lingual gyrus; see Fig. 11B).

Steady-state response in the intermodulation frequency

A, similar to study 1, 4 Hz frequency-tagging activity was increased shortly before target onset when expecting an auditory compared to a visual target (p = .006). D, source localization showed activity over auditory sensory areas, but did not reach significance. N = 27;

Discussion

Earlier version: The neuropsychological account of attention defines it as the selective facilitation (i.e., prioritization) of relevant sensory input and suppression of irrelevant sensory input. Oscillatory activity in the alpha range (∼10 Hz) has been suggested to play a mechanistic role in attention through inhibition of irrelevant cortices, commonly referred to as the ‘alpha inhibition hypothesis’ (Foxe et al., 1998; Jensen & Mazaheri, 2010; Klimesch et al., 2007). In the current cross-modal attention study we directly tested this hypothesis by using frequency-tagging to specifically examine how cues signalling the modality of an upcoming target (either the auditory or visual modality) affected the responsiveness of the relevant and irrelevant sensory cortices prior to target onset. In-line with previous work, we observed a post-cue increase in posterior alpha power in anticipation of processing auditory targets. However, contrary to prevalent theories proposing visual gain suppression when focusing on the auditory modality, we observed that the amplitude of visual frequency-tagging responses increased just prior to the onset of the auditory target. This suggests that responsiveness of the visual stream was not inhibited when attention was directed to auditory processing and was not inhibited by occipital alpha activity, which directly contradicts the alpha inhibition hypothesis. Our results reconcile previously paradoxical results on audio-visual attention and support the view that alpha activity gates downstream communication pathways.

Revised version: The neuropsychological account of attention defines it as the selective facilitation (i.e., prioritization) of relevant sensory input and suppression of irrelevant sensory input. Oscillatory activity in the alpha range (∼10 Hz) has been suggested to play a mechanistic role in attention through inhibition of irrelevant cortices, commonly referred to as the ‘alpha inhibition hypothesis’ (Foxe et al., 1998; Jensen & Mazaheri, 2010; Klimesch et al., 2007). In the current cross-modal attention study we directly tested this hypothesis by using frequency-tagging to specifically examine how cues signalling the modality of an upcoming target (either the auditory or visual modality) affected the responsiveness of the relevant and irrelevant sensory cortices prior to target onset. In-line with previous work, we observed a post-cue increase in posterior alpha power in anticipation of processing auditory targets. However, contrary to prevalent theories proposing visual gain suppression when focusing on the auditory modality, we observed that the amplitude of visual frequency-tagging responses increased just prior to the onset of the auditory target. This suggests that responsiveness of the visual stream was not inhibited when attention was directed to auditory processing and was not inhibited by occipital alpha activity, which directly contradicts the proposed mechanism behind the alpha inhibition hypothesis. Our results reconcile previously paradoxical results on audio-visual attention and support the view that alpha activity gates downstream communication pathways.

Frequency-tagging

Removed in the revised version: In the current experiment, we specifically chose to analyse the cue-to-target interval, where both visual and auditory SSEPs/SSEFs present preparatory states for the upcoming task, independent of task-related processing or performance. The magnitude of auditory SSEPs/SSEFs was increased shortly before target onset when expecting an auditory target compared to a visual target, very much in line with previous reports (e.g., Saupe et al., 2009).

Earlier version: In contrast to the results reported in Saupe et a., (2009), where visual SSEPs decreased when attending the auditory modality, visual SSEPs/SSEFs increased shortly before target onset when expecting an auditory target in our data. This is especially surprising as auditory targets were frequently, or in case of our second study, always accompanied by visual distractors, rendering it optimal for task success to completely ignore any visual input. As auditory targets were significantly more difficult than visual targets in our first study and of comparable difficulty in our second study, these results strongly speak to a vigilance increase of sensory processing independent of modality and an inability to selectively disengage one sensory modality in anticipation of a demanding task. This view is consistent with previous work in which visual SSEPs elicited by irrelevant background stimulation increased with task load in an auditory discrimination task (Jacoby et al., 2012). Furthermore, our results indicate that task demand is a strong candidate to reconcile previously seemingly paradox results, as splitting attention between the auditory and visual system seemed possible in simpler tasks (Driver & Spence, 1998; Saupe et al., 2009) and impossible under high demand (de Jong et al., 2010; Driver, 1996; Driver & Spence, 1998; Spence & Driver, 1996). An alternative account for our findings stems from the evidence, that participants are more likely to only perceive and react to the visual modality, when confronted with audio-visual stimuli (Colavita, 1974; Spence, 2009). However, this effect was mostly limited to speeded modality discrimination/target detection tasks (Sinnett et al., 2008; Spence, 2009). Furthermore, the increased difficulty of the here-used auditory stimuli was confirmed in a previous block-design study (van Diepen & Mazaheri, 2017) and in our second study, performances over the visual and auditory tasks were comparable. Nevertheless, visual dominance could play a role for auditory target difficulty as well as predictions over the reciprocity of the audio-visual relationship.

Revised version: In the current experiment, we specifically chose to analyse the cue-to-target interval, where both visual and auditory SSEPs/SSEFs present preparatory states for the upcoming task, independent of task-related processing or performance. The magnitude of auditory SSEPs/SSEFs was increased shortly before target onset when expecting an auditory target compared to a visual target, very much in line with previous reports (e.g., Saupe et al., 2009). In contrast to the results reported in Saupe et a., (2009), where visual SSEPs decreased when attending the auditory modality, visual SSEPs/SSEFs increased shortly before target onset when expecting an auditory target in our data. This is especially surprising as auditory targets were frequently, or in case of our second study, always accompanied by visual distractors, rendering it optimal for task success to completely ignore any visual input. As auditory targets were significantly more difficult than visual targets in our first study and of comparable difficulty in our second study, these results strongly speak to a vigilance increase of sensory processing independent of modality and an inability to selectively disengage one sensory modality in anticipation of a demanding task. This view is consistent with previous work in which visual SSEPs elicited by irrelevant background stimulation increased with task load in an auditory discrimination task (Jacoby et al., 2012). Furthermore, our results indicate that task demand is a strong candidate to reconcile previously seemingly paradox results, as splitting attention between the auditory and visual system seemed possible in simpler tasks (Driver & Spence, 1998; Saupe et al., 2009) and impossible under high demand (de Jong et al., 2010; Driver, 1996; Driver & Spence, 1998; Spence & Driver, 1996). An alternative account for our findings stems from the evidence, that participants are more likely to only perceive and react to the visual modality, when confronted with audio-visual stimuli (Colavita, 1974; Spence, 2009). However, this effect was mostly limited to speeded modality discrimination/target detection tasks (Sinnett et al., 2008; Spence, 2009). Furthermore, the increased difficulty of the here-used auditory stimuli was confirmed in a previous block-design study (van Diepen & Mazaheri, 2017) and in our second study, performances over the visual and auditory tasks were comparable. Nevertheless, visual dominance could play a role for auditory target difficulty as well as predictions over the reciprocity of the audio-visual relationship.

A revision of the alpha inhibition hypothesis

Earlier version: Top-down cued changes in alpha power have now been widely viewed to play a functional role in directing attention: the processing of irrelevant information is attenuated by increasing alpha power in cortices involved with processing this information (Foxe, Simpson, & Ahlfors, 1998; Hanslmayr et al., 2007; Jensen & Mazaheri, 2010). However, recent evidence suggests that posterior alpha activity does not inhibit gain in early sensory processing stages (Antonov et al., 2020; Gundlach et al., 2020; Gutteling et al., 2022; Zhigalov & Jensen, 2020). To date there has been no direct investigation into the effect of alpha increases on later stages in the processing stream. In the current study, as expected, we observed a post-cue increase in occipital alpha activity in anticipation of an auditory target. However, we also observed an increase in the amplitude of visual SSEPs during the cue-target interval, which directly contradicts the widespread view of alpha activity exerting ‘gain control’ in early sensory areas by regulating excitability (Foxe & Snyder, 2011; Jensen & Mazaheri, 2010; Van Diepen et al., 2019). Here we propose that alpha activity, rather than modulating early primary sensory processing, exhibits its inhibitory effects at later stages of the processing stream (Antonov et al., 2020; Gundlach et al., 2020; Zhigalov & Jensen, 2020; Zumer et al., 2014), gating feedforward or feedback communication between sensory areas (Bauer et al., 2020; Haegens et al., 2015; Uemura et al., 2021). Our data provides evidence in favour of this view, as we can show that early sensory alpha activity covaries over trials with SSEP magnitude in higher order sensory areas. If alpha activity exerted gain control in early visual regions, increased alpha activity would have to lead to a decrease in SSEP responses. In contrast, we observe that increased alpha activity originating from early visual cortex is related to enhanced visual processing. Source localization confirmed that this enhancement was not originating from early visual areas, but from areas associated with later stages of the processing stream such as the precuneus, which has been connected to sensory integration (Al-Ramadhani et al., 2021; Xie et al., 2019). While we cannot completely rule out alternative explanations, it seems plausible to assume that inhibition of other task-irrelevant communication pathways leads to prioritised and thereby enhanced processing over relevant pathways. In line with previous literature (Morrow et al., 2023; Peylo et al., 2021; Zhigalov & Jensen, 2020), we therefore suggest that alpha activity limits task-irrelevant feedforward communication, thereby enhancing processing capabilities in relevant downstream areas (see Fig. 1A). The benefit of alpha-modulation and its effect on visual information processing is underlined by reaction times, which were faster both for trials with high pre-target alpha activity and high pre-target visual SSVEF activity.

Revised version: Top-down cued changes in alpha power have now been widely viewed to play a functional role in directing attention: the processing of irrelevant information is attenuated by increasing alpha power in areas involved with processing this information (Foxe, Simpson, & Ahlfors, 1998; Hanslmayr et al., 2007; Jensen & Mazaheri, 2010). However, recent evidence suggests that posterior alpha activity does not inhibit gain in early sensory processing stages (Antonov et al., 2020; Gundlach et al., 2020; Gutteling et al., 2022; Zhigalov & Jensen, 2020). To date there has been no direct investigation into the effect of alpha increases on later stages in the processing stream. In the current study, as expected, we observed a post-cue increase in occipital alpha activity in anticipation of an auditory target. However, we also observed an increase in the amplitude of visual SSEPs during the cue-target interval, which directly contradicts the widespread view of alpha activity exerting ‘gain control’ in early sensory areas by regulating excitability (Foxe & Snyder, 2011; Jensen & Mazaheri, 2010; Van Diepen et al., 2019). Here we propose that alpha activity, rather than modulating early primary sensory processing, exhibits its inhibitory effects at later stages of the processing stream (Antonov et al., 2020; Gundlach et al., 2020; Zhigalov & Jensen, 2020; Zumer et al., 2014), gating feedforward or feedback communication between sensory areas (Bauer et al., 2020; Haegens et al., 2015; Uemura et al., 2021). Our data provides evidence in favour of this view, as we can show that early sensory alpha activity does not covary over trials with SSEP magnitude in early visual areas, but covaries instead over trials with SSEP magnitude in higher order sensory areas (see also SUPPL. Fig. 14). If alpha activity exerted gain control in early visual regions, increased alpha activity would have to lead to a decrease in SSEP responses. In contrast, we observe that increased alpha activity originating from early visual cortex is related to enhanced visual processing. Source localization confirmed that this enhancement was not originating from early visual areas, but from areas associated with later stages of the processing stream such as the precuneus, which has been connected to sensory integration (Al-Ramadhani et al., 2021; Xie et al., 2019). While we cannot completely rule out alternative explanations, it seems plausible to assume that inhibition of other task-irrelevant communication pathways leads to prioritised and thereby enhanced processing over relevant pathways. In line with previous literature (Morrow et al., 2023; Peylo et al., 2021; Zhigalov & Jensen, 2020), we therefore suggest that alpha activity limits task-irrelevant feedforward communication, thereby enhancing processing capabilities in relevant downstream areas (see Fig. 1A). The benefit of alpha-modulation and its effect on visual information processing is underlined by reaction times, which were faster both for trials with high pre-target alpha activity and high pre-target visual SSVEF activity.

It should be noted, the comparison between modulation in alpha activity and in SSEP/SSEFs is difficult, especially concerning timing. This is largely owed to differences in signal-to-noise due to trial averaging in the frequency versus the time domain and temporal and frequency lag in the estimation of alpha activity (Peylo et al., 2021). It is further noteworthy, that the majority of evidence for the alpha inhibition hypothesis focused on the effect of pre-target alpha modulation on behaviour and target-related potentials (Morrow et al., 2023). However, in our data alpha modulation occurs clearly ahead of SSVEP/SSVEF modulation on a scale that could not be simply explained by temporal or frequency smearing. Additionally, significant trial-by-trial correlations, which occur in the frequency domain for both signal types, underline the strong relationship between both measurements.

Earlier version: Interestingly, we could show that the magnitude of the correlation between alpha power and visual information processing varied between conditions, suggesting a dynamic and adaptive regime. This notion supports the view that alpha oscillations represent a mechanism rather than a specific function, which can fulfil different roles depending on task demand and network location, which has been confirmed in a recent study revealing functionally distinct alpha networks (Clausner et al., 2024). As such, it is conceivable that alpha oscillations can in some cases inhibit local processing, while in other cases, depending on network location, connectivity and demand, alpha oscillation can facilitate signal transmission. In different contexts, utilizing unimodal targets and distractors, spatial cueing, or covert attention, different functional processes could be involved (Morrow et al., 2023). Future research should intensify efforts to disentangle these effects, investigating localized alpha networks intracranially or through combinations of fMRI, EEG and MEG, to clearly measure their effects on sensory processing and behaviour.

Revised version: Interestingly, we could show that the magnitude of the correlation between alpha power and visual information processing varied between conditions, suggesting a dynamic and adaptive regime. This notion supports the view that alpha oscillations represent a mechanism rather than a specific function, which can fulfil different roles depending on task demand and network location, which has been confirmed in a recent study revealing functionally distinct alpha networks (Clausner et al., 2024). As such, it is conceivable that alpha oscillations can in some cases inhibit local transmission, while in other cases, depending on network location, connectivity and demand, alpha oscillation can facilitate signal transmission. This mechanism allows to increase transmission of relevant information and to block transmission of distractors.

Added in the revised version: In different contexts, utilizing unimodal targets and distractors, spatial cueing, or covert attention, different functional processes could be involved (Morrow et al., 2023). Future research should intensify efforts to disentangle these effects, investigating localized alpha networks intracranially or through combinations of fMRI, EEG and MEG, to clearly measure their effects on sensory processing and behaviour.

It is known that the localisation of alpha activity reflects the retinotopic organisation of visual spatial attention topographically over the parietooccipital cortex (Kelly et al., 2006; Popov et al., 2019). Notably, recent studies provided evidence, that the same organisation can be observed for auditory attention. Specifically, the localisation of visual alpha activity in the parietooccipital cortex reflects the spatial direction of auditory attention (Klatt et al., 2021; Popov et al., 2021). This observation can be explained through micro-saccades towards the spatial location of sounds, which are irrevocably connected to alpha oscillations (Popov et al., 2021). While we did not manipulate spatial attention, our results fit well to the notion of visual alpha activity serving as a sensory orientation system, relaying visual information to task-relevant downstream processing areas, and blocking communication to irrelevant pathways.

The intermodulation frequency

Previous research showed that simultaneous frequency-tagging in multiple frequencies evokes a response in the intermodulation frequency (f1 – f2). In multimodal settings, this frequency is thought to reflect cross-modal integration (Drijvers et al., 2021). This is very well in line with our findings, where increased vigilance of the sensory system arising from anticipation of a difficult auditory target resulted in an increase in the intermodulation frequency. Likewise, our data shows that visual signal enhancement was localized in precuneus, which has been connected to sensory integration (Al-Ramadhani et al., 2021; Xie et al., 2019). Furthermore, we could show that the intermodulation frequency covaries over trials with reaction time in the easiest condition, where visual targets were presented without any distractors. A lack of this connection in other conditions might reflect increasing interferences from higher task difficulty, rather than a lack of the effect itself, but this remains to be tested. We cannot exclude an alternative explanation, as theta oscillations are known to be involved in movement preparation, it is possible that phase-resets could lead to time-locked appearance of these oscillations (Lakatos et al., 2008; Tomassini et al., 2017).

Conclusion

Our results taken together suggest that under high attention, audio-visual excitability is enhanced, reflecting an increase in vigilance for both visual as well as auditory information, even if this increases processing of distracting information. We showed that this vigilance shift, as reflected by SSEP/SSEF responses, is regulated by early visual alpha activity, presumably through relaying of visual information over communication pathways, thereby controlling the downstream flow of visual information.

Materials and Methods

Participants EEG-Study

In total, 24 healthy volunteers participated in this study (mean age: 19.1 ± 1.8 SD; 17 women). Due to technical difficulties, one participant could not finish the experiment and one participant did not exceed chance level in the behavioural task (∼33 %). Both were therefore removed from any further analysis. All remaining participants reported normal or corrected-to-normal vision, no history of psychiatric or neurological illness and provided written informed consent. After completion of the experiment, participants received either monetary compensation or certification of their participation for their university course program. The study protocol was approved by the Ethics Committee of the School of Psychology at the University of Birmingham and is in accordance with the Declaration of Helsinki.

Participants MEG-Study

In total, 28 healthy volunteers participated in this study (mean age: 23.4 ± 3.6 SD; 20 women). One participant was removed from further analysis, as they only responded to ∼42% of trials correctly in the second block, which related to 27/19 trials per condition respectively. All remaining participants reported normal or corrected-to-normal vision, no history of psychiatric or neurological illness and provided written informed consent. After completion of the experiment, participants received either monetary compensation or certification of their participation for their university course program. The study protocol was approved by the Ethics Committee of the School of Psychology at the University of Birmingham and is in accordance with the Declaration of Helsinki.

Cross-modal attention paradigm EEG-Study

Each trial was initiated by a brief presentation of a cue (100 ms) signalling the modality of the upcoming discrimination task (v-shape: visual modality; inversed v-shape: auditory modality; diamond-shape: non-specific). During the following cue-to-target interval, the fixation cross was frequency-tagged at 36 Hz. At the same time, a 40 Hz frequency-tagged sound (amplitude modulated white noise) was played over headphones (see Fig. 1). The volume of tones was initially adjusted to a level that was clearly perceivable but not uncomfortable and remained stable over participants. No task was connected to this interval and participants did not need to pay attention to either the sound or the fixation cross. After three seconds, the frequency-tagging stopped, and right after the cessation of stimuli, the target was presented for a very brief moment (25 ms). It consisted either of a Gabor patch (visual modality) or a sound (auditory modality). If the target modality was visual, participants had to use the three arrow buttons on the keyboard to indicate whether the Gabor patch was tilted to the left (−10°; left arrow button), vertical (0°; down arrow button), or tilted to the right (10°; right arrow button). Additionally, in 50% of the trials, a random distractor from the pool of auditory targets was presented simultaneously to the visual target over headphones. Similarly, if the target modality was auditory, participants used the same buttons to indicate whether the pitch of the tone was low (500 Hz), medium (1000 Hz) or high (2000 Hz). Again, in 50% of the trials, a random distractor from the pool of visual targets was also presented. If the cue was non-specific (diamond shape), either a visual (50% of non-specific trials) or an auditory target was presented, never accompanied by any distractors. Experimental trials were separated by an inter-trial interval of 4 seconds, to avoid carry-over effects from previous trials. The resulting 6 conditions were randomly ordered and balanced out over the experiment.

During the experiment, participants were instructed to keep their gaze locked to a fixation cross presented at the centre of the screen. Preceding data collection, participants performed 36 practice trials to get accustomed to the task and the target stimuli. The ensuing experiment was split into 26 trial sequences, separated by self-chosen breaks, which together resulted in 468 trials and lasted between 80 and 90 minutes. The discrimination task was programmed and presented with MATLAB® R2020b and Psychtoolbox-3 on an LCD-monitor featuring a 140 Hz refresh rate. The onset of the visual and auditory tagging frequencies (i.e steady state stimuli) were tracked using the Cedrus Stimtracker (https://cedrus.com/stimtracker/index.htm).

Adjustments to the attention paradigm in the MEG-Study

In our second study, we removed all ambiguity concerning targets and distractors and therefore developed a blocked design, incorporating two blocks. The first block did not display distractors and only cues which correctly predicted the target were presented. Cues in the second block were likewise always correctly indicating the target modality, but this time, each target was accompanied by a random distractor from the non-target modality.

Furthermore, the visual task was adjusted to be more in line with the auditory task. As such, the Gabor patches now featured stripes in different frequencies (e.g. a low number of stripes, a medium number of stripes and a high number of stripes. The participant’s task was to discriminate between these three Gabor patches. As auditory targets had been markedly more difficult in our first study, we now included a brief difficulty calibration prior to the experiment. First, we presented 21 Gabor patches with 3 different amounts of stripes following a standard difficulty. If participants could discriminate them correctly 55 – 75% of the time, this difficulty setting was chosen. Otherwise, depending on the performance, the stripe-frequency of the Gabor patches was adjusted. There were maximally 3 sessions of 21 Gabor patches, after which we had enough data to calibrate the individual difficulty setting.

The same procedure was then performed with the difficulty of the tones, calibrating the pitch frequency for each individual participant.

Lastly, visual frequency-tagging stimulation now followed a sinusoidal contrast-change rather than an on-off stimulation, which was possible due to a high-resolution projector featuring a refresh rate of 1440 Hz (PROPixx DLP LED projector ;VPixx Technologies Inc., Canada).

Eye-tracking

To confirm that participants had focused on the fixation cross during the cue-to-target interval, we incorporated eye-tracking into our MEG-experiment (EyeLink 1000 Plus). Correct trials of the second block were analysed for vertical and horizontal eye-movements. To exclude blinks from this analysis, trials with very large eye-movements (> 10 degrees of visual angle) were removed from the eye-tracking data (See suppl Fig. 5).

Behavioral analysis

We were interested in accuracy in the discrimination of visual targets and auditory targets, as well as reaction times. Furthermore, we examined the distraction cost of having a target presented with a distractor of a different modality as well as the reaction time to make the target discrimination. The distraction cost was calculated as the reaction time difference between cued targets with distractors (i.e. visual and auditory stimuli presented together) and cued targets without distractors (either a visual or auditory stimulus presented alone). All incorrect trials as well as trials with reaction times faster than 100 ms or exceeding 1500 ms were removed from analysis (0.5 % too fast, 9.2% too slow).

EEG data acquisition

All EEG recordings were conducted using a WaveGuard Cap (ANTneuro), featuring 64 Ag/AgCL electrodes (10-10 system; ground: Fz; reference: Cpz; EOG: left canthus). Electrodes positions were prepared with OneStep cleargel conductive paste and impedances were kept below 100 kΩ. The measured signal was transmitted using an ANTneuro EEGosports amplifier (low-pass filter: 150 Hz; high-pass filter: 0.5 Hz; sampling rate: 500 Hz).

MEG data acquisition

Prior to the experiment, feducial positions and head-shape were recorded using a FASTRAK system (Polhemus, USA). The experiment took place in a dimly lit room, where participants were seated in a comfortable chair in the gantry of a 306-sensor TRIUX Elekta system with 204 orthogonal planar gradiometers and 102 magnetometers (Elekta, Finland). The 71*40 cm screen was positioned at ∼1.40 m distance from the participant.

EEG Preprocessing

Offline analyses were performed in MATLAB ® R2020b. The data was pruned from artifacts by visual inspection using the EEGLAB toolbox (Delorme & Makeig, 2004). Additionally, blinks and ocular artefacts were removed from the data using independent component analysis (ICA). EEG channels were re-referenced to an average of all channels (excluding EOG).

MEG Preprocessing

Offline analyses were performed in MATLAB ® R2020b and Python. Spatiotemporal Signal-Source-Separation (SSS) was applied to the raw data via MNE’s inbuilt maxfilter function with a duration window of 10 s and a correlation value of .9. The data was pruned from artifacts by visual inspection using the Fieldtrip toolbox (Delorme & Makeig, 2004). Additionally, blinks and ocular artefacts were removed from the data using independent component analysis (ICA). In sensor space, planar gradiometers were combined for further analyses. In source space, all individual planar gradiometers were analysed.

Amplitude of the evoked frequency-tagging response

To investigate the temporal dynamics of amplitude of the frequency tagged responses after the onset of the attentional cues (also the precise onset of the frequency tagged stimuli) the data was epoched into 6-second segments starting 1.5 seconds prior to cue onset. Next the data were narrow-band filtered around the 36 Hz activity to capture the visual frequency-tagging, 40 Hz activity in to capture the auditory frequency-tagging, and the intermodulation frequency at 4 Hz, which can be derived by subtracting both frequency-tagging responses (fi = fauditory – fvisual; see Drijvers, Spaak & Jensen, 2020). Here we used a Blackmann-windowed sync filters adapted to a suitable ratio of temporal and frequency resolution for the specific frequency of each of the tagged signals: filter order 116 for 35.5 to 36.5 Hz and 39.5 to 40.5, filter order 344 for 3.5 to 4.5 Hz. The filtered data at each of the tagged frequencies as well the intermodulated frequency were baseline corrected (interval between 700 and 200 ms preceding cue onset) before calculating the average over trials to obtain steady-state evoked potentials). The power envelope of the SSEPs of tagged frequencies was estimated using Hilbert transformation. To confirm that 4 Hz is a sufficient distance between tagging frequencies, we repeated to analysis for 43.5 to 44.5. We found no indication of frequency-bleeding over, as the effects observed at 40 Hz, were not present at 44 Hz (see SUPPL Fig. 11).

Temporal dynamics of the induced EEG changes

In addition to looking at the cue evoked changes in the amplitude of the frequency tagged signals, we investigated the induced changes in the EEG signal at the frequencies of the tagged auditory and visual stimuli, alpha activity (9-11 Hz, filter order 276 for 7.5 to 12.5 Hz), as well as the intermodulation frequency (4 Hz; filter order 344 for 3.5 to 4.5 Hz). Here rather than averaging the epoched data filtered at the specific frequency ranges, we performed the Hilbert transform, and averaged the power-envelope of the specific frequencies across trials. This approach is very much analogous to the standard time-frequency analysis using convolutions (van Diepen & Mazaheri, 2017; Zhigalov et al., 2019), but affords more control concerning temporal versus frequency resolution to examine the temporal dynamics of the specific frequencies of interest. In our second study, the individual peak alpha frequency was used for bandpass-filtering in contrast to a standardised band applied in the first study.

Time-frequency representations of power

In addition to the estimating frequency power envelopes, Time–frequency representations (TFRs) of power of the EEG signal were estimated using the Fieldtrip toolbox (Oostenveld et al., 2011). The power or frequencies between 5 and 20 Hz were calculated for each trial, using a sliding time window (frequency steps: 0.5; time steps: 10 ms). The length of the window was adjusted to a length of 3 cycles per frequency and tapered with a Hanning window. For each trial, both datasets were normalised to display relative percent change from baseline using the following formula: [(activity – baseline) / baseline], where baseline refers to the interval between 700 and 200 ms before cue onset. To estimate the topographical distribution of voltage differences between conditions, uncorrected power values were normalised applying the following formula: [Δ/∑ = (a – b) / (a + b)], where a and b reflect the different conditions.

Added in the revised version: Source localization

Earlier version: Source localizationSource localization was performed with a beamformer approach using the Fieldtrip toolbox. Headmodels were created based on individual T1-scans fitted to fiducial points and head shapes. These data were fit to a 5mm 3d sourcemodel and warped into MNI-space. Two participants were missing individual T1-scans. In these cases, we applied a standardized T1-scan using the Colin 27 Average Brain Model (Holmes et al., 1998). Frequency-domain data was localized using the Dynamic Imaging of Coherent Sources (DICS beamformer) method with a dpss taper of 2 Hz assuming fixed orientation. As condition differences between frequency-tagging responses were better estimated in the time-domain, we assessed them applying the Synthetic Aperture Magnetometry (SAM-beamformer) method with optimal fixed rotation (Sekihara et al., 2004). Significant brain areas and peak coordinates were related to brain areas using the Anatomical Automatic Labeling (AAL) atlas for SPM8 (Tzourio-Mazoyer et al., 2002) . After statistical analysis, source localized data was interpolated onto the Colin 27 Average Brain Model MRI. Cerebellar and brainstem interpolations were excluded from the coordinate system.

Revised version: Source localization was performed with a beamformer approach using the Fieldtrip toolbox. Headmodels were created based on individual T1-scans fitted to fiducial points and head shapes. These data were fit to a 5mm 3d sourcemodel and warped into MNI-space. Two participants were missing individual T1-scans. In these cases, we applied a standardized T1-scan using the Colin 27 Average Brain Model (Holmes et al., 1998). Frequency-domain data was localized using the Dynamic Imaging of Coherent Sources (DICS beamformer) method with a dpss taper of 2 Hz assuming fixed orientation. As condition differences between frequency-tagging responses were better estimated in the time-domain, we assessed them applying the Synthetic Aperture Magnetometry (SAM-beamformer) method with optimal fixed rotation (Sekihara et al., 2004). Significant brain areas and peak coordinates were related to brain areas using the Anatomical Automatic Labeling (AAL) atlas for SPM8 (Tzourio-Mazoyer et al., 2002) . After statistical analysis, source localized data was interpolated onto the Colin 27 Average Brain Model MRI. Cerebellar and brainstem interpolations were excluded from the coordinate system.

Statistical analysis

Condition differences in the behavioural task were estimated with paired t-tests and repeated measures ANOVAS, utilizing Tukey-Kramer Post-Hoc test. Power differences between conditions, as well as source-space contrasts were analysed using cluster permutation analysis (Maris & Oostenveld, 2007). In this procedure, condition labels were randomly shuffled 1000 times, creating pairs of surrogate conditions. To test for significance, paired t-tests were conducted for each data point and each channel, resulting in one t-matrix for real conditions and 1000 t-matrices for surrogate conditions. Significant t-values (p < .05) were defined as clusters if there was at least one significant data point present at the same time and frequency in at least two neighbouring channels. To correct for multiple comparisons, a condition difference was only assumed, if the maximum sum of t-values in a real cluster exceeded the same sum of 95% of the clusters found in the surrogate data. To replicate our results in the second study, we applied the same statistics and averaged the previously found time-intervals into windows of 500ms (e.g.: if we found an effect −0.51 to −0.0620 s prior to target onset we tested this effect for a time window of −0.5 to 0 s prior to target onset).

The relationship between induced changes in alpha activity and frequency tagged responses was assessed using trial by trial Spearman correlations. For each participant and each electrode, a correlation coefficient was calculated between the average activity in a previously identified cluster, which was used as seed (e.g., condition differences in alpha activity), and the average activity of the electrophysiological correlate of interest (e.g., 36 Hz activity over the previously identified time window). The correlation coefficients were z-transformed, and the resulting channel by participant matrix was tested against null-correlation model using the cluster permutation approach described above. Derived clusters were additionally tested and visualised by comparing median split trials of high vs low activity. For this analysis, outliers (values deviating more than 2 standard deviations from the mean) were excluded. Furthermore, the average correlation coefficient of the cluster was tested against a 0-correlation model for each participant using t-statistics. Lastly, an interaction between electrophysiological correlations and conditions was performed using correlation coefficients for each participant and electrode, testing them between conditions using the cluster permutation approach. Perceptually uniform and universally readable colormaps were applied to all visualisations (Crameri et al., 2020). All data are presented as mean ± standard error of the mean (SEM).